---
title: 'Data Loading'
description: 'Import CSV Data'
---

## Example coding exercise

```yaml
type: NormalExercise
key: 2bafef99a3
lang: r
xp: 100
skills: 1
```

This is an example exercise. The aim is to load data and analyze times series. 

`@instructions`


`@hint`


`@pre_exercise_code`
```{r}

```

`@sample_code`
```{r}

```

`@solution`
```{r}
ls()

```

`@sct`
```{r}

```
